---
  title: "one"
  description: "one"

---
## New test

```yaml
type: NormalExercise

xp: 100

key: c2820290ac



```

Add number

`@instructions`
Add numbers

`@hint`
come on

`@pre_exercise_code`
```{undefined}
a = 1
b = 2
```

`@solution`
```{undefined}
3
```
`@sct`
```{undefined}
a + b
```



