---
  title: "example_python_challenges"
  output: "html_document"
  v2: false

---
## Choose the right answer

```yaml
type: MultipleChoiceChallenge

xp: 0

key: 115qFYuQ4pPlifqZze0QkDORLwVufanGIfC



```














---
## concatenating lists

```yaml
type: BlanksChallenge

xp: 0

key: 8a565a9166



```














---
## popping lists (2)

```yaml
type: OutputChallenge

xp: 0

key: 6258197972



```














---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 100

key: 7eebf45376



```

Yo, try to add numbers

`@instructions`
What i said above

`@hint`
come on

`@pre_exercise_code`
```{undefined}
a = 1
b = 2
```

`@solution`
```{undefined}
3
```
`@sct`
```{undefined}
a + b
```



